# CL5: Loops

Welcome to the fifth coding lab!

In this Coding Lab we will focus on debugging and understanding loops.

**Note**: You may *not* be able to complete all of this in 50 minutes, particularly the challenges in Part 3.  This are meant to help you gain mastery and provide additional practices, not to stress you out a ton. Make a concerted effort. Do your best to complete it. But, if you don't figure it all out, that's why we post answers for coding labs!

## Part 1: Debugging

### Creating Errors
For each of the follow Exceptions, write a small piece of code that will cause it to happen:
- ZeroDivisionError
- SyntaxError
- IndentationError
- TypeError
- IndexError
- ValueError
- NameError

<!-- grade_id: cell-0e441a93d9074f56, points: 0, grade: true, solution: true -->
```python
# ZeroDividisionError
one = 1
zero = 0

one / zero
```

Output:
```
ZeroDivisionError                         Traceback (most recent call last)
/tmp/ipykernel_157/671190287.py in <module>
      3 zero = 0
      4 
----> 5 one / zero

ZeroDivisionError: division by zero
```

<!-- grade_id: cell-b1ac7106e83cb72f, points: 0, grade: true, solution: true -->
```python
# SyntaxError
if true:
    print('hello world!')
```

Output:
```
NameError                                 Traceback (most recent call last)
/tmp/ipykernel_157/1855108694.py in <module>
      1 # SyntaxError
----> 2 if true:
      3     print('hello world!')

NameError: name 'true' is not defined
```

<!-- grade_id: cell-cfb3db08e0f1e2f2, points: 0, grade: true, solution: true -->
```python
# IndentationError
if zero == 0:
print('your number is zero !')
```

<!-- grade_id: cell-ffce1ddb5088e527, points: 0, grade: true, solution: true -->
```python
# TypeError
winter_quarter = 'UCSD' + 2022
```

<!-- grade_id: cell-425c5dbe3ce96cb2, points: 0, grade: true, solution: true -->
```python
# IndexError
info = ['Sasha', 'Zabegalin', '6/14/2000', 'UCSD', 'Cogs Sci: ML']
info[5]
```

Output:
```
IndexError                                Traceback (most recent call last)
/tmp/ipykernel_157/3175564383.py in <module>
      1 # IndexError
      2 info = ['Sasha', 'Zabegalin', '6/14/2000', 'UCSD', 'Cogs Sci: ML']
----> 3 info[5]

IndexError: list index out of range
```

<!-- grade_id: cell-16b2a47d0a4214d2, points: 0, grade: true, solution: true -->
```python
# ValueError
int('number')
```

Output:
```
ValueError                                Traceback (most recent call last)
/tmp/ipykernel_157/1961010247.py in <module>
      1 # ValueError
----> 2 int('number')

ValueError: invalid literal for int() with base 10: 'number'
```

<!-- grade_id: cell-6438bfb83cd2bd54, points: 0, grade: true, solution: true -->
```python
# NameError
name = 'Sasha'
nme 
```

Output:
```
NameError                                 Traceback (most recent call last)
/tmp/ipykernel_157/3994713475.py in <module>
      1 # NameError
      2 name = 'Sasha'
----> 3 nme

NameError: name 'nme' is not defined
```

### Fixing Errors

For each of the following cells, try to debug the code to fix any errors, and get it running. 

<!-- grade_id: cell-5c1e4912f8be4917, points: 0, grade: true, solution: true -->
```python
if True:
    print('It worked')
```

Output:
```
It worked
```

<!-- grade_id: cell-a311736fa20712bf, points: 0, grade: true, solution: true -->
```python
age_string = "My favorite number is " + str(27)

print(age_string)
```

Output:
```
My favorite number is 27
```

<!-- grade_id: cell-49b6889e712f45f6, points: 0, grade: true, solution: true -->
```python
print('Three plus five equals = ' + str((3+5)))
```

Output:
```
Three plus five equals = 8
```

## Part 2: Loops Review & Exploration

We will now explore two ways to iterate over data in Python - `for` and `while` loops. 

Each type of loop performs its respective code block until certain conditions are met. Here, we will explore what those conditions are. 

Note: "iterate" means to perform an action repeatedly 

Note: Sometimes, we accidentally make loops that never end. 
- If that happens, "break" the cell by selecting it and then pressing "ctrl+c". Alternatively, press the "stop" button (square icon) in the Jupyter toolbar. 

### Loops 

First, declare a variable called `counter` that has type int and value zero.

Then write a `while` loop that runs 10 times and use `counter` to track how many times the loop runs. 

<!-- grade_id: cell-af48110fe42b669b, points: 0, grade: false, solution: true -->
```python
counter = 0

while counter < 10:
    counter += 1
    print(counter)
```

Output:
```
1
2
3
4
5
6
7
8
9
10
```

```python
# Check that (after the loop runs) that there is a variable called `counter` that has value 10
assert counter == 10
```

### Infinite Loops

The while loop below will run indefinitely unless someone puts a stop to it! 

Add your own code that causes the loop to exit when the `counter` variable reaches 10. You can use `break` to do this. 

Note: Sometimes, we accidentally make loops that never end - try that by running the cell below before adding 

If that happens, interrupt the cell by selecting it and then pressing "ctrl+c". Alternatively, press the "stop" button (square icon) in the Jupyter toolbar. 

For refernece, the starter code provided is:
```python
counter = 0

while True:
    
    print("the value of 'counter' = " + str(counter))
    counter +=1
    
```

<!-- grade_id: cell-0e33a3d206c8759d, points: 0, grade: false, solution: true -->
```python
counter = 0

while True:
    
    print("the value of 'counter' = " + str(counter))
    counter +=1
    
    if counter == 10:
        break
```

Output:
```
the value of 'counter' = 0
the value of 'counter' = 1
the value of 'counter' = 2
the value of 'counter' = 3
the value of 'counter' = 4
the value of 'counter' = 5
the value of 'counter' = 6
the value of 'counter' = 7
the value of 'counter' = 8
the value of 'counter' = 9
```

Now, rather than using `break`, edit the loop from above so that the condition following `while` will eventually evaluate as `False` (and thus terminate the loop).

<!-- grade_id: cell-b561c8f0d73a25e3, points: 0, grade: false, solution: true -->
```python
counter = 0

while counter < 10:
    counter +=1
```

### Loop Explorations

Now, explore using for and while loops and breaks. 

Here are some things to try, and questions to ask - each of which you should explore in code. 
- What happens if you use 'break' all by itself?
- What are some different ways to loop through a list? tuple?
- Write these combinations: 
    - a `for` loop with a `break`, a `for` loop with a `continue`
    - a `while` loop with a `break`, a `while` loop with a `continue`
- Write some loops that loop across lists:
    - Write a loop that loops across, and prints out, all the elements of a list
    - Write a loop that loops across, and prints out, every second element of a list
    - Write a loop that loops across, and prints out, every element of a list in reverse order
    - Write a loop that loops across, and prints out, the first half of a list
- Write a loop with a conditional inside it.
- Can you have loops inside loops? If you think so, try and write a nested loop. 

Reminder: you can add as many cells as you'd like to test these out! 

<!-- grade_id: cell-c101bd32c5150b72, points: 0, grade: false, solution: true -->
```python
break
```

Output:
```
File "/tmp/ipykernel_157/668683560.py", line 1
    break
    ^
SyntaxError: 'break' outside loop
```

```python
tup = (1, 2, 3, 4, 5)

for thing in tup:
    print(thing)
```

Output:
```
1
2
3
4
5
```

```python
index = 0
lst = ['this', 'is', 'a', 'list']

while index < len(lst):
    
    print(lst[index])
    index += 1
```

Output:
```
this
is
a
list
```

```python
for item in lst:
    if item == 'a':
        break
    
    print(item)
```

Output:
```
this
is
```

```python
for item in lst:
    if item == 'a':
        continue
    
    print(item)
```

Output:
```
this
is
list
```

```python
index = 0

while index < len(lst):
    print(lst[index])
    index += 1
    
    if lst[index] == 'a':
        break
```

Output:
```
this
is
```

```python
index = 0

while index < len(lst):
    
    if lst[index] == 'a':
        continue
    
    print(lst[index])
    index += 1
```

Output:
```
this
is
```

### Specific looping task

`print`ing the output is great to learn about and explore loops, but typically we're looping to generate some output. 

Generate a list below that stores various integers and strings. 

Write a loop that will separate out the elements in your list by type, such that all the integers are stored in a list called `my_ints` and all of the strings are stored in a list called `my_strs`

<!-- grade_id: cell-77fd056f3c67c2a3, points: 0, grade: false, solution: true -->
```python
my_lst = ['Sasha', 'Zabegalin', 6, 14, 2000]
my_ints = []
my_str = []

for item in my_lst:
    if type(item) == str:
        my_str.append(item)
    elif type(item) == int:
        my_ints.append(item)
        
print(my_ints)
print(my_str)
```

<!-- grade_id: cell-362fd3c5dd6ea1ca, points: 0, grade: true, solution: false -->
```python
assert type(my_ints) == list
assert type(my_strs) == list

# check to make sure all values in each list are of expected type
assert all([type(x) == int for x in my_ints])
assert all([type(x) == str for x in my_strs])
```

## Part 3: Challenges

Give tese a shot, but it's definitely ok if you don't figure these out all the way! Remember, if you don't figure out the answer, definitely check the answer key when posted!

### Debugging Challenge

You're trying to write code that will count the number of vowels in your name, storing that value in `counter`. To get started, you store your name in the variable `my_name` and write the code you see below. However, the code below is not quite working...work to debug the code below to accomplish the task!

Note: This question uses a `for` loop. Next week's coding lab will get you lots of loops practice!

<!-- grade_id: cell-1d48166464642c57, points: 0, grade: false, solution: true -->
```python
my_name = 'Sasha Zabegalin'
vowels = ['A','E','I','O','U', 'a', 'e', 'i', 'o', 'u']
counter = 0

for char in my_name:
    if char in vowels:
        counter +=  1

print('|', my_name + '\'s vowel total:', counter, '|')
```

### Nested Loop Challenge

Using nested loops, you will print out a pattern that looks like this:

1 <br>
1 2 <br>
1 2 3 <br>
1 2 3 4 <br>
1 2 3 4 5 <br>
1 2 3 4 5 6 <br>
1 2 3 4 5 6 7

Hints:
- You can do this with two while loops, one inside the other.
- You will have two indices counters, one for each loop. 
- For the print statement, it will look something like `print(index_inner, end=" ")`

<!-- grade_id: cell-3be5b519bda3c9fd, points: 0, grade: false, solution: true -->
```python
def pyrmaid(nums):
    
    for row in range(len(nums)):
        for col in range(0, row+1):
            print(nums[col], end=' ')
        print('\n')

nums = ['1', '2', '3', '4', '5', '6', '7', '8', '9']
pyrmaid(nums)


```

If you get the above pattern working, try manipulating the code to print out different patterns, and with different symbols. 

<!-- grade_id: cell-d7ff87fc6930bed0, points: 0, grade: true, solution: true -->
```python
stars = ['*', '*', '*', '*', '*']
pyrmaid(stars)
```

## The End!

This is the end of this Coding Lab!

Be sure you've made a concerted effort to complete all the tasks specified in this lab. Then, go ahead and submit on datahub!